Before you turn this assignment in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says "YOUR ANSWER HERE" or `YOUR CODE HERE` and remove the `raise NotImplementedError()` lines. 

Code blocks starting with a `# tests ...` comment provide unit tests which have to run without
errors in order to get full points. Be aware that there might be further 'secret' tests to check correct
implementation! I.e. the provided unit tests are necessary but not sufficient for full points!


Please add your name and student ID below:

```python
NAME = "Fabian Traxler"
STUDENT_ID = "01553958"
```

<!-- grade_id: student_id_test, points: 0, grade: true, solution: false -->
```python
assert len(NAME) > 0
assert len(STUDENT_ID) > 0
```


# Intelligent Audio and Music Analysis Assignment 0

The main purpose of this assignment is to get you familiarized with the infrastructure and workflow.  

## Task 1: Import numpy

Load the python math library [numpy](https://numpy.org/doc/stable/) into namespace variable `np`.


<!-- grade_id: import-numpy, points: 0, grade: false, solution: true -->
```python
# load the python math library numpy into a local namespace variable called "np"
# this is achieved using the 
#
# import <lib> as <name>
#
# syntax

# YOUR CODE HERE
import numpy as np
```

<!-- grade_id: import-numpy-test, points: 1, grade: true, solution: false -->
```python
# test if numpy has been loaded successfully

np.version.version

print('All tests successful!')
```

Output:
```
All tests successful!
```

## Task 2: Create a Matrix

Create a matrix ([ndarray](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.html)) with size 4x4 which contains random values and store it in a variable named `x`.

<!-- grade_id: create-matrix, points: 0, grade: false, solution: true -->
```python
# numpy has a module called "random", use the rand function. 
# See documentation at: 
# https://numpy.org/doc/stable/reference/random/generated/numpy.random.rand.html

# YOUR CODE HERE
x = np.random.rand(4,4)

print("x = {}".format(x))
```

Output:
```
x = [[0.35009395 0.59017319 0.70232656 0.53517294]
 [0.6695761  0.98766938 0.23680238 0.481775  ]
 [0.33253566 0.26225061 0.15223324 0.01792852]
 [0.14920659 0.52574414 0.0244324  0.47055614]]
```

<!-- grade_id: create-matrix-test, points: 1, grade: true, solution: false -->
```python
# test x for properties:

assert x.shape == (4, 4)
assert all([0 <= ex <= 1 for ex in x.flatten()])

print('All tests successful!')
```

Output:
```
All tests successful!
```

<!-- grade_id: create-matrix-exception-test, points: 1, grade: true, solution: false -->
```python
# test for exception: 

try:
    x[10, 10]
except IndexError:
    pass
else:       
    raise AssertionError('did not raise ValueError')
    
print('All tests successful!')
```

Output:
```
All tests successful!
```

## Task 3: Elementwise Matrix Operations
Calculate the *elementwise* square of `x` into a new variable named `y`.

<!-- grade_id: multiply-matrix, points: 0, grade: false, solution: true -->
```python
# Arrays can be used as matrices, thus we can choose between matrix multiplication and elementwise multiplication.
# https://numpy.org/doc/stable/reference/generated/numpy.matmul.html
# https://numpy.org/doc/stable/reference/generated/numpy.multiply.html

# YOUR CODE HERE
y = np.multiply(x,x)
```

<!-- grade_id: cell-01f26ea77ab64c5e, points: 1, grade: true, solution: false -->
```python
# Test shape of y
# Note: for full points, values of y must also be correct.

assert y.shape == (4, 4)


print('All tests successful!')
```

Output:
```
All tests successful!
```

## Task 4: Freeform Answers

Sometimes you will find questions that require a textual answer. Use the markdown input field below to submit your answer.
 
Answer the following question:

**What is your background or experience with programming, audio processing, and machine learning?**


<!-- grade_id: cell-33dfd1f4d9e317f6, points: 1, grade: true, solution: true -->
Started programming during my Bachelor Program and have done some private projects as well as some projects during Internships. I have no prior Knowledge in audio processing. I first familiarized myself with machine learning in my Bachelor Program and have since gained way more experience during an internship and the Master Programm Data Science.

```python

```